## <span style="margin-left: 1px; background-color: #333; color: white; padding:8px; padding-left: 10px; padding-right: 10px; border-radius: 5px;">ECC201 • Fall 2020-2021 • <span style="margin-left: 1px; background-color: #cf441e; border-radius: 3px; padding: 2px; padding-left: 5px; padding-right: 5px; font-size: 80%;">QUIZ 02</span></span>

#### <span style="display: block; margin-left: 0px; background-color: #388; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 0px; text-align: center; border-radius: 3px;">1 question / 30 points total / 30 minutes </span></div><span id="QTOP"></span> 


Jump to 
<a href="#Q1" style="background-color: #555; color: white; padding: 10px; border-radius: 2px; text-align: center; border-radius: 3px; margin-right: 5px;">Q1</a>

### INSTRUCTIONS <span style="margin-left: 10px; background-color: #116; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 5px; font-size: 12pt;">READ CAREFULLY</span>

</br></br>

<span style="color: #555; font-size: 12pt; line-height: 1.5;">
    
    
- All questions in this quiz/exam will be automatically graded using [NBGRADER](https://nbgrader.readthedocs.io/).
    
- <div style="background-color: #800; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; margin-top: 10px; margin-bottom: 10px; border-radius: 7px;">
        
    Your answers will be checked for potential cheating. If there is clear evidence of cheating, then you will automatically fail this quiz/exam. If you use online websites or other means to receive any sort of "help," you are committing **CHEATING**, and it will **not** be tolerated! Even if you cheat on a single question or one part of a question, you will still receive a **zero** for the entire exam, and further disciplinary action may also be taken. So, **please do not cheat!**
        
    </div>
    
    
- The maximum amount of time that you should be spending on each question or section is provided to you in a box colored <span style="margin-left: 0px; background-color: #dd6a00; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 5px;">orange</span> like this.
    
    
- **<span style="color: #080; font-size: 12pt; line-height: 1.5;">USE *Jupyter Notebook* to solve this exam!</span> <span style="color: #b00; font-size: 12pt; line-height: 1.5;">DO NOT USE *Jupyter Lab*!</span>**

    
- __<span style="color: #b00; font-size: 12pt; line-height: 1.5;">DO NOT COPY AND PASTE CELLS</span>__

    
- __<span style="color: #b00; font-size: 12pt; line-height: 1.5;">DO NOT REMOVE CELLS</span>__
    
    
- __<span style="color: #b00; font-size: 12pt; line-height: 1.5;">DO NOT ADD NEW CELLS</span>__    

    
- __<span style="color: #b00; font-size: 12pt; line-height: 1.5;">DO NOT CHANGE THE TYPE (Markdown/Code/Raw) OF CELLS</span>__    

    
- __<span style="color: #080; font-size: 12pt; line-height: 1.5;">ONLY WRITE CODE IN DESIGNATED CELLS</span>__    

    
- This is an exam just like any in-class exam you are used to taking until now. You will only be taking this exam at your convience in your home.

    
- Fill in the code cells indicated for you clearly. That is, do __not__ write your code anywhere else. If you do so, you may lose the points for that question.
    - You will see
    
            # YOUR CODE HERE
            pass
           
      in the functions where you will have to fill in your Python code. Replace these with your own code so that your implementation works as specified.    
    
    
- Do not change function names (and argument names/lists). Otherwise, this may cause you to lose points on a question unnecessarily.

- If you cannot finish all the questions on this exam, submit as much as you managed to complete so that you can get some of the points available for that exam.
    
    
- Submit the completed version of this notebook at the web address indicated below.
- Do not send a file with a different file type. Only Jupyter notebooks will be accepted.
    
    
- Do not send your answer notebook file as email or via any other messaging media&mdash;such requests will be ignored.
    
    
- Do not ask any questions about the exam unless you believe there is something wrong with a question. Remember this is an exam.
    
    
- If you fail to upload your file before the deadline for any reason whatsoever, this is entirely your responsibility. In such cases, your grade for the exam in question will be zero.
    
    
- __All _Rules of Academic Conduct_ apply__. Any sort of cheating will result in an __FF__ grade for the course. Your case will also be submitted to the Faculty Disciplinary Committee for further evaluation and disciplinary action. You can refer to books and web pages for help, but you are __not__ allowed to consult anyone online. Any sort of collaboraton with anyone anywhere is considered cheating.
    - See https://neu.edu.tr/administration/regulations/general-student-discipline-regulations/?lang=en
    
- Enter your official name and your NEU student ID in the code cell below. For example,

        student_name = "Huseyin Sevay"
        student_id   = "20201083"
        
    - Make sure you enter your official name as a `str` as shown above
    - Make sure that your student ID has been entered correctly as `str`.


- ### SUBMISSION
    
    - Rename your file so that that its name contains your official name and student ID. For example,
    
            huseyin_sevay_20201083_ecc201-fall-2020-2021-quiz01.ipynb

    - Upload your Jupyter notebook with your answers at the appropriate UZEM link.
        
    - Check the exact deadline date and time before you start this test.

    - No late submissions will be accepted under any circumstances!
  
    
- <span style="background-color: #800; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; margin-top: 10px; margin-bottom: 10px; border-radius: 7px;"> WARNING</span> __<span style="color: #800; line-height: 1.5;">Failure to follow any of these instructions may cost you points on this exam!</span>__

### Best of luck!
</span>

### Your name and student ID <span style="margin-left: 10px; background-color: #116; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 5px; font-size: 12pt;">Enter in the code cell below</span>

```python
student_name = "Özlem Ezgi SARI"
student_id   = "20171403"
```

---

### <span id="Q1" style="background-color: #555; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 5px;">Question 1</span><span style="margin-left: 10px; background-color: #383; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 5px;">30 points</span><span style="margin-left: 10px; background-color: #dd6a00; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 5px;">30 mins</span><a href="#QTOP" style="background-color: #aaa; color: white; padding: 5px; padding-bottom: 5px; border-radius: 2px; text-align: center; border-radius: 3px; margin-left: 10px; text-decoration: none;">⇡</a> 

<div style="margin-top: 1px;">&nbsp;</div>


Implement a Python *iterator class* named `RandomSequenceIterator` that returns the elements of a given sequence in *random* index order. The iterator must take, as first parameter the input sequence, and, as second parameter, a randomizer function, which randomizes the order of the input sequence indices. 

Inspect the `__init__()` method provided for you already.

In addition, see the examples in the test cases further below.

## <span style="margin-left: 10px; background-color: #811; color: white; padding: 5px; margin: 10px; border-radius: 5px; font-size: 12pt;">YOU MUST WRITE ALL YOUR CODE IN THE FOLLOWING CODE CELL</span>##

<!-- grade_id: cell-a8i39e0k88740a8g, points: 0, grade: false, solution: true -->
```python
class RandomSequenceIterator:
    def __init__( self, seq, randomizer ):
        self.seq = seq
        self.rindices = randomizer( range( len( seq ) ))
        #print( f"--- self.rindices: {self.rindices}" )
        self.index = 0
        self.newSeq = list()


    def __iter__(self):
        return self

    def __next__(self):
        for i in self.rindices:
            seq = self.seq[i]
            self.newSeq.append(seq)
            
        if self.index < len(self.newSeq):
            newSeq = self.newSeq[self.index]
            self.index += 1
            return self.newSeq
        else:
            raise StopIteration
#raise NotImplementedError()
```

```python
# TEST 1
#
import random

random.seed( 100 )

L = ["A", "B", "C", "D", "E"]
randomizer = lambda s: random.sample( s, len( s ) )
R = RandomSequenceIterator( L, randomizer )
myiter = iter(R)
print(next(myiter))
correct_answer = ['B', 'D', 'E', 'A', 'C']

assert list( R ) == correct_answer
```

**TEST CASES:**

<!-- grade_id: cell-q5559su24fybp676, points: 10, grade: true, solution: false -->
```python
# TEST 1
#
import random

random.seed( 100 )

L = ["A", "B", "C", "D", "E"]
randomizer = lambda s: random.sample( s, len( s ) )
R = RandomSequenceIterator( L, randomizer )
correct_answer = ['B', 'D', 'E', 'A', 'C']

assert list( R ) == correct_answer

```

<!-- grade_id: cell-11j2h35aqpfnd7gn, points: 10, grade: true, solution: false -->
```python
# TEST 2
#
import random

random.seed( 200 )

S = "Oregano"
randomizer = lambda s: random.sample( s, len( s ) )
R = RandomSequenceIterator( S, randomizer )
correct_answer = ['O', 'r', 'o', 'n', 'e', 'g', 'a']

assert list( R ) == correct_answer

```

<!-- grade_id: cell-9agu0il13xd14f8d, points: 10, grade: true, solution: false -->
```python
# TEST 3
#
import random

random.seed( 500 )

S = "Peace and Justice"
randomizer = lambda s: random.sample( s, len( s ) )
R = RandomSequenceIterator( S, randomizer )

first = next( R )

#print( first )

assert first == 'i'

rest = [v for v in R]

#print( rest )

correct_answer = ['e', 'e', 'J', 's', 'u', 'a', 'c', 'e', ' ', 't', 'c', 'P', 'n', ' ', 'd', 'a']

assert rest == correct_answer

```

---

<span style="display: inline-block; width: 100%; background-color: #dd6a00; color: white; padding: 5px; padding-left: 10px; padding-right: 10px; border-radius: 5px; text-align: center; font-weight: bold;">END OF TEST</span>